Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` (you should of course delete `raise NotImplementedError()` which is there only as reminder), while not modifying the other cells (but you should run them to check the output you obtain). Please also fill your first and last name and your VU ID below:

```python
STUDENT_FIRST_NAME = "Begüm" # e.g. "John" (no "J", no "J.", no "John S.M.")
STUDENT_LAST_NAME = "Yalçın"  # e.g. "Smith"
VU_ID_NUM = "2797623"          # e.g. "2789012"
```

---

# Assignment 1

The goal of this assignment is to make you familiar with Python's syntax, data types, and file I/O. This assignment is divided into 9 exercises, which are worth different amount of points (total is 100 points).

We assume that the folder that you work in is the one obtained by unzipping the given ``assignment1.zip`` file and thus has the following structure. When working on your submission, please respect it, otherwise the autograder will give errors.

<code>
assignment1.ipynb
data/...
</code>


## Important remarks
- **Working together**: You are meant to work individually on the first three assignments. You can, of course, brainstorm ideas and discuss issues with your fellow students, but you are required to write your solutions individually. 
- **Plagiarism**: All your code will be automatically scanned for plagiarism. Furthermore, using the internet as a passive resource is allowed. This means that you can search for help there and partially copy code, as long as you explicitly acknowledge inside your Jupyter notebook which parts have been copied and from where. 
- **Performance**: You should optimize the computational performance of your functions. Specifically, when grading the assignments, we set a hard limit of 15 seconds for each cell execution. This should be sufficient to cover each of the test cases. Function calls that take longer than 15 seconds will not be awarded any points.
- **Code styling**: Your implementation will not be checked for style. However, we do encourage you to practice good code styling. See, for example, https://docs.python-guide.org/writing/style/.
- **Chronological run**: All outputs should be repeatable by doing one full “chronological” run of the notebook without any manual changes to code blocks, including parameters. (try it yourself by clicking ``Kernel -> Restart and run all``, which should give the result as handed in).
- **Handing in**: Hand in the .ipynb file of your notebook on the [Canvas page for Assignment 1](https://canvas.vu.nl/courses/68046/assignments/262961) by *Sunday June 11th 23:59*. 
- **Other questions**: If you have doubts/questions about the assignments, feel free to ask them in [this discussion thread](https://canvas.vu.nl/courses/68046/discussion_topics/648148) so that everyone will be able to see them and our answers. 

```python
# Remember to run this cell before any of the following ones

import random
import numpy as np
from numpy.testing import assert_equal
```

## Exercise 1: Type specific operations (5 points)

Write a function that performs operations depending on specific types of the input.
- `bool`: return the negation
- `int`: square the number and subtract 1
- `float`: round the number to three decimal places
- `complex`: return the difference between the real part and the imaginary part (for example, given the number $5 + 3 \mathrm{i}$, the answer should be $5-3=2$)
- `string`: return the palindrome string whose first half is the given string, e.g., 'abc' becomes 'abccba'
- `list`: move the first element to the last, e.g., `[1, 2, 3]` -> `[2, 3, 1]`
- `tuple`: replace the first element of the tuple with the last element, e.g., `(1, 'b', [3])` -> `([3], 'b', [3])`
- `set`: remove all elements that are not integers or floats, e.g., `{'1', 2, 'c', 4.222, (5, 67)}` -> `{2, 4.222}`
- `dict`: reverse all (key, value) pairs. If the value is not hashable, then remove the (key, value) pair. You may assume that all keys and values are unique.
- For any other type, return `None`.

Be aware that your function will be tested with inputs of length up to 50000 items.

<!-- grade_id: Q1-code, points: 0, grade: false, solution: true -->
```python
def operate(item):
    if isinstance(item, bool):
        return not item
    elif isinstance(item, int):
        return item**2 - 1
    elif isinstance(item, float):
        return round(item, 3)
    elif isinstance(item, complex):
        return item.real - item.imag
    elif isinstance(item, str):
        return item+ item[::-1]
    elif isinstance(item, list):
        return item[1:] +[item[0]]
    elif isinstance(item, tuple):
        return (item[-1],) + item[1:]
    elif isinstance(item, set):
        res = set()
        for i in item:
            if isinstance(i, (int, float)):
                res.add(i)
        return res
    elif isinstance(item, dict):
        new_dict = {}
        for key, value in item.items():
            try:
                hash(value)
                new_dict[value] =key       
            except TypeError:
                pass
        return new_dict

    else:
        return None

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: Q1-test, points: 5, grade: true, solution: false -->
```python
# Test your function on some possible inputs

assert_equal(operate(True), False)
assert_equal(operate(99), 99**2-1)
assert_equal(operate(0.987654321), 0.988)
assert_equal(operate(0.9), 0.9)
assert_equal(operate(complex(10, 4)), 6)
assert_equal(operate("abc"), "abccba")
assert_equal(operate([1, 2, 3]), [2, 3, 1])
assert_equal(operate((1, "b", [3])), ([3], "b", [3]))
assert_equal(operate({"1", 2, "c", 4.222, (5, 67)}), {2, 4.222})
assert_equal(operate({"a": 1, "b": [2], "c": (3,)}), {(3,): "c", 1: "a"})

```

## Exercise 2: Coin flips streaks (5 points)

Write a Python function that takes as input a single string of any length describing a sequence of coin flips with possible outcome head (``'H'``) or tail (``'T'``) and returns only the length of the longest coin flip streak. 

The function should return always an integer and, in particular, should return ``0`` if the string is empty. You can check the correctness of your function using some instances of strings provided below with the corresponding correct answers. 

Your function will be tested with string consisting of up to 1 million coin flips. To generate yourself longer strings of length ``l`` with probability of head equal to ``p`` (and tail then equal to ``1-p``), you can use the command 

```coin_flips = ''.join(random.choices(['H','T'], weights = [p,1-p], k = l))```

If you do so, remember to use the command ```random.seed(number)``` to get reproducible results.

<!-- grade_id: Q2-code, points: 0, grade: false, solution: true -->
```python
def find_longest_streak(coin_flips):
    longest = 1
    current = 1
    #check if empty
    if len(coin_flips)==0:
        return 0
    for i in range(len(coin_flips)-1):
        if coin_flips[i]== coin_flips[i-1]:
            current=current+1
            if current> longest:
                longest = current #update longest
        else:
            current=1 # reset current 
        
    return longest
    
#raise NotImplementedError()
```

<!-- grade_id: Q2-test, points: 5, grade: true, solution: false -->
```python
# Test your function on some possible coin flip sequences

assert_equal(find_longest_streak(""), 0)
assert_equal(find_longest_streak("T"), 1)
assert_equal(find_longest_streak("HTTHT"), 2)
assert_equal(find_longest_streak("TTHTTTHHTHHH"), 3)
assert_equal(find_longest_streak("HTHHTTTHTHTHHHTTHHHHHHTHTTTHHTHHHTHTHTHTTTHHHHHTHTHTTH"),6)

```

## Exercise 3: parsing unique words in a text (5 points)

Write a function that takes a .txt file as input and returns the list of unique words contained in that text. Note that:

- Ignore capitalized letters, namely if a words appears both capitalized (e.g., ``'Until'``) and not capitalized (e.g., ``'until'``), it should appear only once in the final unique word list;
- The text might appear in various paragraphs separated by new lines and/or have empty lines;
- The text has punctuation, which your function should ignore. Specifically, you should ignore all (and only) the punctuation returned by ``string.punctuation``
- Related to the previous point, you can assume the text has **no abbreviations** (e.g., *don't*), no **English possessive** (e.g., *John's*) and no **hypened words** (e.g., *full-scale*)
- Numbers should be processed according to the above rules (regardless of whether they appear in time/date/percentages/decimals). For instance, processing the text ``On June 7th, at 9:00am`` should return the unique words ``['on','june','7th','at','900am']``.

Note that you can use the Python `string` module (imported below), but *no* other packages for this exercise.

<!-- grade_id: Q3-code, points: 0, grade: false, solution: true -->
```python
import string

def unique_words_parser(filename):
    set_words = set()
    try:
        with open(filename, "r") as f:           
            content= f.read()
            # remove punctuation
            for punctuation in string.punctuation:
                content = content.replace(punctuation, "")
            words = content.split()
            for word in words:
                word= word.lower() #lowercase
                word = word.strip() #strip 
                set_words.add(word)
        return list(set_words) #convert set to list 
    except FileNotFoundError:
        print("File not found.")
        return None 

#raise NotImplementedError()
```

<!-- grade_id: Q4-test, points: 5, grade: true, solution: false -->
```python
# Test your function on some possible input files. Note that only the length of the returned word list is checked.
# The autograder, however, will check the exact word list.

assert_equal(len(unique_words_parser("data/text_input0.txt")), 85)
assert_equal(len(unique_words_parser("data/text_input1.txt")), 149)

```

## Exercise 4: Scrabble (15 points)

### Part A (5 points)
Given a list ``letters`` of (possibly repeated) available letters, write a function should return a list with:
- *the longest word* that can be formed using the letters in the list ``letters`` that is a valid Scrabble word, and
- *its length* (that is, the number of its letters).

If there are multiple words with the same length, the function returns the **last** one in alphabetical order. If there are no words that can be formed using the letters in the list, the function returns the empty string and its length. The list of all admissible Scrabble words is provided in a csv file whose ``filename`` should also be passed as first argument. 

Your function will be tested with list of letters of size up to 50 and on different dictionary .csv files, not necessarily alphabetically sorted. Your function should not edit the input .csv file nor create any other file. The provided tests all use the file ``data/dictionary.csv``, which contains all admissible English Scrabble words.

You can use the package ``csv`` (imported below) but you cannot import any other package.

<!-- grade_id: Q4a-code, points: 0, grade: false, solution: true -->
```python
import csv

def find_longest_word(filename, letters):
    longest_word = ""
    length = 0

    with open(filename, 'r') as f:
        csv_file = csv.reader(f)
        for row in csv_file:
            is_valid = True
            word = row[0]   
            for letter in word:
                # if a word occurs more than letters it is not valid 
                if word.count(letter) >letters.count(letter):
                    is_valid = False
                    break
            if is_valid:
                # check if word şs longer than longest
                if len(word) > length or (len(word) == length and word > longest_word):
                    longest_word = word
                    length = len(word)
    return longest_word, length
    
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: Q4a-test, points: 5, grade: true, solution: false -->
```python
# Test your function on a few possible inputs

filename = "data/dictionary.csv"
assert_equal(find_longest_word(filename, []), ("", 0))
assert_equal(find_longest_word(filename, ["a", "a", "m", "m", "m", "n", "w", "i", "n"]), ("manna", 5))
assert_equal(find_longest_word(filename, ["q", "z", "i", "u", "k", "l", "e", "a", "y"]), ('quaylike', 8))
assert_equal(find_longest_word(filename, ["x", "l", "f", "d", "n", "z", "w"]), ("", 0))
assert_equal(find_longest_word(filename, ["w", "e", "a", "e", "t", "h", "l", "z"]), ('wealth', 6))
assert_equal(find_longest_word(filename, ["q", "l", "d", "f", "c", "o", "a", "l", "m", "n", "w", "l", "o", "o", "n"]), ("monoclonal", 10))

```

### Part B (10 points)

Given a list ``letters`` of (possibly repeated) available letters, write a function should return a list with:
- *the most valuable word* that can be formed using the letters in the list ``letters`` that is a valid Scrabble word, and
- *its value*.

The value of a word is equal to the sum of the values of its letters. The value of each letter is given by a dictionary ``letter_values``, which should be passed as second argument to the function. 

If there are multiple words with the same value, the function returns the **first** one in alphabetical order. If there are no words that can be formed using the letters in the list, the function returns the empty string and its length. The list of all admissible Scrabble words is provided in a csv file whose ``filename`` should also be passed as first argument. 

Your function will be tested with list of letters of size up to 20 and on different dictionary .csv files, not necessarily alphabetically sorted. You can assume the words in the dictionary are all lowercase. Your function should not edit the input .csv file nor create any other file. You can use the packages ``csv`` but you cannot import any other package.

<!-- grade_id: Q4b-code, points: 0, grade: false, solution: true -->
```python
def find_most_valuable_word(filename, letters, letter_values):
    most_valuable_word = ""
    valuable_value = 0
    with open(filename, 'r') as f:
        csv_file = csv.reader(f)
        for row in csv_file:
            word = row[0]
            # check if word can be formed 
            if len(word) <= len(letters):
                is_valid = True
                for letter in word:
                    # if a word occurs more than letters it is not valid 
                    if word.count(letter) >letters.count(letter):
                        is_valid = False
                        break
                if is_valid:
                    word_value = 0
                    for letter in word:
                        # sum values of letters 
                        word_value += letter_values.get(letter, 0)
                    # check if word şs more valuable  than most valuable
                    if word_value > valuable_value or (word < most_valuable_word and word_value == valuable_value):
                        most_valuable_word = word
                        valuable_value = word_value

    return most_valuable_word, valuable_value
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: Q4b-test, points: 10, grade: true, solution: false -->
```python
# Test your function on a few possible inputs

filename = "data/dictionary.csv"
letter_values = {
    "a": 1, "b": 3, "c": 3, "d": 2, "e": 1, "f": 4, "g": 2, "h": 4,
    "i": 1, "j": 8, "k": 5, "l": 1, "m": 3, "n": 1, "o": 1, "p": 3,
    "q": 10, "r": 1, "s": 1, "t": 1, "u": 1, "v": 4, "w": 4, "x": 8,
    "y": 4, "z": 10
}

assert_equal(find_most_valuable_word(filename, [], letter_values), ("", 0))
assert_equal(find_most_valuable_word(filename, ["a", "a", "m", "m", "m", "n", "w", "i", "n"], letter_values), ("mamma", 11))
assert_equal(find_most_valuable_word(filename, ["q", "z", "i", "u", "k", "l", "e", "a", "y"], letter_values), ('queazy', 27))
assert_equal(find_most_valuable_word(filename, ["x", "l", "f", "d", "n", "z", "w"], letter_values), ("", 0))
assert_equal(find_most_valuable_word(filename, ["w", "e", "a", "e", "t", "h", "l", "z"], letter_values), ('hazel', 17))
assert_equal(find_most_valuable_word(filename, ["q", "l", "d", "f", "c", "o", "a", "l", "m", "n", "w", "l", "o", "o", "n"], letter_values), ('downfall', 15))

```

## Exercise 5: Nasty numbers (10 points)

A positive integer is said to be **nasty** if it has a prime number of ones in its binary representation. Write a Python function with two arguments, two nonnegative integers ``a`` and ``b``, the second of which has to be optional.
- If only ``a`` is provided, the function should return the list of all nasty numbers between ``0`` and ``a`` (included).
- If both ``a`` and ``b`` are provided, the function should return the list of all nasty numbers between them (both included), regardless if $a\leq b$ or $a>b$.

Note that you can use the Python `math` module (imported below), but *no* other packages for this exercise.

<!-- grade_id: Q5-code, points: 0, grade: false, solution: true -->
```python
import math

def check_prime(n):
    if n<2:
        return False
    if n==2:
        return True
    root = int(math.isqrt(n))
    for i in range(2, root+1):
        if n%i==0:
            return False
    return True

def nasty_numbers(a, b = None):
    nasty_numbers = []
    if b is None:
        first = 0
        last = a
    else:
        first = min(a,b)
        last = max(a,b)
    for i in range(first, last+ 1):
        # convert to binary representation
        binary = bin(i)[2:]
        # count occurences of binary 
        cnt = binary.count("1")
        # check if count is prime 
        if check_prime(cnt):
            nasty_numbers.append(i)

    return nasty_numbers
   
#raise NotImplementedError()
```

<!-- grade_id: Q5-test, points: 10, grade: true, solution: false -->
```python
# Test your function on some possible input files. 

assert_equal(nasty_numbers(4,11), [5, 6, 7, 9, 10, 11])
assert_equal(nasty_numbers(9), [3, 5, 6, 7, 9])
assert_equal(nasty_numbers(23,14), [14, 17, 18, 19, 20, 21, 22])

```

## Exercise 6: Light choreography (15 points)

For a light art festival, you decide to create a choreography with one million light bulbs on a $1000 \times 1000$ grid.

The rows and columns are both numbered from 0 to 999 and each light bulb is uniquely identified by a pair $r,c$, where $r$ is the row and $c$ is column number where the light bulb lies. In particular the light bulbs at each of the four corner are at 0,0, 0,999, 999,999, and 999,0. 

### Part A (5 points)

The choreography starts with all light bulbs turned off and consists of a set of consecutive instructions describing whether to ``turn on``, ``turn off``, or ``toggle`` various inclusive ranges given as coordinate pairs. Each coordinate pair represents opposite corners of a rectangle, inclusive; a coordinate pair like 0,0 through 2,2 therefore refers to 9 light bulbs in a 3x3 square. 

For example:
- ``turn on 0,0 through 999,999`` would turn on (or leave on) every light bulb.
- ``toggle 0,0 through 999,0`` would toggle the first line of 1000 light bulbs, turning off the ones that were on, and turning on the ones that were off.
- ``turn off 499,499 through 500,500`` would turn off (or leave off) the middle four light bulbs.

Write a Python function that takes as input a .txt file with a list of instructions (with only a single instruction per row) and returns the number of light bulbs that are turned on after executing all the instructions.

Note that you can use the Python `numpy` module (imported above), but *no* other packages for this exercise.

<!-- grade_id: Q6a-code, points: 0, grade: false, solution: true -->
```python
import numpy as np

def do_instruction(line, grid):
    splitted = line.strip().split()
    # get the mode 
    mode = splitted[0] if splitted[0] != 'turn' else splitted[1]
    # get start and end from splitted parts
    start, _, end = splitted[-3:]
    firstx, firsty = map(int, start.split(','))
    lastx, lasty = map(int, end.split(','))
    
    # mode cases 
    if mode== 'on':
        grid[firstx:lastx + 1, firsty:lasty+1] =True
    elif mode == 'off':
        grid[firstx:lastx + 1, firsty:lasty+1] = False
    elif mode == 'toggle':
        grid[firstx:lastx + 1, firsty:lasty+1] = np.logical_not(grid[firstx:lastx + 1, firsty:lasty+1])
    return grid

def artlight_partA(filename):    
    grid = np.zeros((1000, 1000), dtype=bool)
    with open(filename, 'r') as f:
        for line in f:
            grid = do_instruction(line, grid)
    return np.sum(grid)

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: Q6a-test, points: 5, grade: true, solution: false -->
```python
# Test your function on some possible input files. 

assert_equal(artlight_partA("data/lights_0.txt"), 134)
assert_equal(artlight_partA("data/lights_1.txt"), 741352)
assert_equal(artlight_partA("data/lights_2.txt"), 265609)

```

### Part B (10 points)

You got the choreography instructions wrong. Each of the light bulb has individual brightness controls, with the brightness being either zero or a positive integer. At the begining of the choreography, the light bulbs all start at zero brightness.

- The command ``turn on`` actually means that you should increase the brightness of those light bulbs by $1$.
- The command ``turn off`` actually means that you should decrease the brightness of those light bulbs by $1$, to a minimum of $0$.
- The command ``toggle`` actually means that you should increase the brightness of those light bulbs by $2$.

For example:

- ``turn on 0,0 through 0,0`` would increase the total brightness by $1$.
- ``toggle 0,0 through 999,999`` would increase the total brightness by $2000000$.

Write a Python function that takes as input a .txt file with a list of instructions (with only a single instruction per row) and returns the total brightness of all light bulbs after executing all the instructions.

<!-- grade_id: Q6b-code, points: 0, grade: false, solution: true -->
```python
def do_instruction(instructions, grid):
    for instruction in instructions:
        splitted = instruction.strip().split()
        first_splitted = splitted[0]
        mode = splitted[1]    
        # check if first part toggle or turn 
        if first_splitted == 'turn':
            # start and end coordinates 
            firstx, firsty = map(int, splitted[2].split(','))
            lastx, lasty = map(int, splitted[4].split(','))
            # apply on or off
            if mode == 'on':
                grid[firstx:lastx + 1, firsty:lasty+1] += 1
            elif mode == 'off':
                grid[firstx:lastx + 1, firsty:lasty+1] = np.maximum(grid[firstx:lastx + 1, firsty:lasty+1] - 1, 0)
        elif first_splitted == 'toggle':
            # start and end coordinates 
            firstx, firsty = map(int, splitted[1].split(','))
            lastx, lasty = map(int, splitted[3].split(','))
            
            grid[firstx:lastx + 1, firsty:lasty+1] += 2
    return grid

def artlight_partB(filename):
    grid = np.zeros((1000, 1000), dtype=int)  
    with open(filename, 'r') as f:
        instructions = f.readlines()
        grid= do_instruction(instructions, grid)
    return np.sum(grid)

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: Q6b-test, points: 10, grade: true, solution: false -->
```python
# Test your function on some possible input files.

assert_equal(artlight_partB("data/lights_0.txt"), 247)
assert_equal(artlight_partB("data/lights_1.txt"), 2320807)
assert_equal(artlight_partB("data/lights_2.txt"), 1313035)

```

## Exercise 7: Bingo (15 points)

### Part A (5 points)

Bingo is played on a set of boards each consisting of a `5x5` grid of numbers. Numbers are chosen at random between `0` and `99` (including), and the chosen number is marked on all boards on which it appears. (Numbers may not appear on all boards.) If all numbers in any row or any column of a board are marked, that board wins. (Diagonals don't count.)

The file in input contains first a sequence of randomly drawn numbers (without repetition) and a random set of boards. For example:

```
7,4,9,5,11,17,23,2,0,14,21,24,10,16,13,6,15,25,12,22,18,20,8,19,3,26,1

22 13 17 11  0
 8  2 23  4 24
21  9 14 16  7
 6 10  3 18  5
 1 12 20 15 19

 3 15  0  2 22
 9 18 13 17  5
19  8  7 25 23
20 11 10 24  4
14 21 16 12  6

14 21 17 24  4
10 16 15  9 19
18  8 23 26 20
22 11 13  6  5
 2  0 12  3  7
```

After the first eleven numbers are drawn (7, 4, 9, 5, 11, 17, 23, 2, 0, 14, and 21), there are no winners. Finally, 24 is drawn and the third board wins because it has at least one complete row or column of marked numbers (in this case, the entire top row is marked: 14 21 17 24 4).

The score of the winning board can now be calculated. Start by finding the sum of all unmarked numbers on that board; in this case, the sum is ``188``. Then, multiply that sum by the number that was just called when the board won, ``24``, to get the final score, ``188 * 24 = 4512``.

Write the function ``bingo_first`` that finds which board wins first given a specific input file and outputs only the corresponding score (as an ``int``) calculated as detailed above.

Note that you can use the Python `numpy` module (imported above), but *no* other packages for this exercise.

<!-- grade_id: Q7a-code, points: 0, grade: false, solution: true -->
```python
def bingo_first(filename):
      with open(filename, 'r') as f:
        content = f.readlines()
        line = content.pop(0)
        boards_array = get_boards(content)
        winning_board_index = -1
        num_array = np.array(list(map(int, line.strip("").split(','))))
        score = None
        for num in num_array: 
            indices = np.where(boards_array == num)
            if len(indices) > 0:
                is_found = False
                # get indices of numbers 
                for i in range(len(indices[0])):
                    x = indices[0][i]
                    y = indices[1][i]
                    z = indices[2][i] 
                    boards_array[x, y, z] = -1
                    # checks if a row or column has all -1 
                    if np.all(boards_array[x, :, z] == -1):
                        winning_board_index = x
                        is_found = True
                        break
                    elif np.all(boards_array[x, y, :] == -1):
                        winning_board_index = x
                        is_found = True
                        break
                    else:
                        isFound = False
                if is_found:
                    break       
        if winning_board_index != -1:
            unmarked_sum = boards_array[x][boards_array[x] > 0].sum()
            score = unmarked_sum * num  
        return score

def get_boards(content):
    boards_list = []
    board = np.zeros((5, 5), dtype=int)
    for i, line in enumerate(content):
        lines = []
        for x in line.strip("").split():
            lines.append(int(x))
        # check if it is last line of board 
        if i % 6== 5:  
            loc =(i % 6) - 1
            board[loc] = lines
            boards_list.append(board)
            board = np.zeros((5, 5), dtype=int)
        elif i %6 !=0:
            loc =(i % 6) - 1
            board[loc] = lines
    return np.array(boards_list)
    
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: Q7a-test, points: 5, grade: true, solution: false -->
```python
# Test your function with any one of the following test input files

assert_equal(bingo_first('data/bingo_input0.txt'), 79789)
assert_equal(bingo_first('data/bingo_input1.txt'), 64780)
assert_equal(bingo_first('data/bingo_input2.txt'), 63424)

```

### Part B (10 points)

We are now interested in finding out which board would win **last**. In the above example, the second board is the last to win, which happens after 13 is eventually called and its middle column is completely marked. If you were to keep playing until this point, the second board would have a sum of unmarked numbers equal to ``148`` for a final score of ``148 * 13 = 1924``.

Write the function ``bingo_last`` that finds which board wins last for a given input file and outputs only the corresponding score (as an ``int``) calculated as detailed above.

<!-- grade_id: Q7b-code, points: 0, grade: false, solution: true -->
```python
def get_boards(content):
    boards_list = []
    board = np.zeros((5, 5), dtype=int)
    for i, line in enumerate(content):
        lines = []
        for x in line.strip("").split():
            lines.append(int(x))
        if i % 6== 5:  
            loc =(i % 6) - 1
            board[loc] = lines
            boards_list.append(board)
            board = np.zeros((5, 5), dtype=int)
        # check if it is last line of board 

        elif i %6 !=0:
            loc =(i % 6) - 1
            board[loc] = lines
    return np.array(boards_list)

def bingo_last(filename):
    with open(filename, 'r') as f:
        content = f.readlines()
    line = content.pop(0)
    boards_array = get_boards(content)
    winning_board_index = -1
    num_array = np.array(list(map(int, line.strip("").split(','))))
    game_array = np.full(100, False, dtype=bool)
    score = None
    for num in num_array:
        indices = np.where(boards_array == num)
        if len(indices) > 0:
            is_found = False
            # get indices of numbers 
            for i in range(len(indices[0])):
                x = indices[0][i]
                y = indices[1][i]
                z = indices[2][i] 
                boards_array[x, y, z] = -1
                # checks if a row or column has all -1 
                if np.all(boards_array[x, :, z] == -1):
                    game_array[x] = True
                    if np.sum(game_array == False)== 0:                        
                        winning_board_index = x
                        is_found = True
                        break
                elif np.all(boards_array[x, y, :] == -1):
                    game_array[x] = True
                    if np.sum(game_array == False)== 0:                        
                        winning_board_index = x
                        is_found = True
                        break
                else:
                    isFound = False
            if is_found:
                break
    if winning_board_index != -1:
        unmarked_sum = boards_array[winning_board_index][boards_array[winning_board_index] > 0].sum()
        score = unmarked_sum * num       
    return score
    
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: Q7b-test, points: 10, grade: true, solution: false -->
```python
# Test your function with any one of the following test input files

assert_equal(bingo_last('data/bingo_input0.txt'),1300)
assert_equal(bingo_last('data/bingo_input1.txt'),35224)
assert_equal(bingo_last('data/bingo_input2.txt'),23541)

```

## Exercise 8: Bill Splitting (15 points) 

A restaurant bill is given as a text file, where each line corresponds to a dish, the price, and the names of people that took part in eating that dish. For example:
```
Pizza Napolitana, 12.00, Bob, Alice
Mac and Cheese, 9.00, Bob, Charly
```
The goal of this exercise is to write a Python function that splits the bill.

**Assumptions**:
- Dishes may occur more than once on the bill.
- For each dish and corresponding names on a line, all people take equal share in eating.
- All prices are given as numbers with two decimals. Use floats to represent the numbers.

### Part A (5 points)
Write a function that reads the bill file and returns a list of lists. Each inner list contains as first element the dish name, as second element the price, and as remaining elements the names of people eating from the dish. Assuming that we are given the example bill above as text file, then your function should return
```
[['Pizza Napolitana', 12.0, 'Bob', 'Alice'], ['Mac and Cheese', 9.0, 'Bob', 'Charly']]
```

<!-- grade_id: Q8a-code, points: 0, grade: false, solution: true -->
```python
def parse_bill(filename):
    bill = []
    with open(filename, 'r') as f:
        for line in f:
            dish, price, *names_people = line.strip().split(', ')
            bill.append([dish, float(price), *names_people])
    return bill
    
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: Q8a-test, points: 5, grade: true, solution: false -->
```python
# Test your function with the provided test input files

assert_equal(parse_bill("data/billsplit_input0.txt"), [["Etna DOC", 18.22, "Michael"], ["Crescia", 11.6, "Michael"]])

# These tests only verify the length; feel free to modify it accordingly.

assert_equal(len(parse_bill("data/billsplit_input0.txt")), 2)
assert_equal(len(parse_bill("data/billsplit_input1.txt")), 5)
assert_equal(len(parse_bill("data/billsplit_input2.txt")), 8)
assert_equal(len(parse_bill("data/billsplit_input3.txt")), 10)
assert_equal(len(parse_bill("data/billsplit_input4.txt")), 1000)

```

### Part B (10 points)
Write a function that reads the bill file and returns a dictionary. The dictionary keys are names and the corresponding value should indicate the money that they have to pay. Make sure that all values are rounded down to two decimal numbers. 

We recommend you to make use of the parsing function that you wrote in **part A**. Be aware that your function will be tested on vary large test files.

**Example**: Assume that you are given the bill shown earlier:
```
Pizza Napolitana, 12.00, Bob, Alice
Mac and Cheese, 9.00, Bob, Charly
```

Your function should output the following:
```
{'Bob': 10.5, 'Alice': 6.0, 'Charly': 4.5}
```


<!-- grade_id: Q8b-code, points: 0, grade: false, solution: true -->
```python
def split_bill(filename):
    bill = parse_bill(filename)
    bill_dict = {}   
    for dish in bill:
        dish_name = dish[0]
        price = dish[1]
        num_people = len(dish) - 2
        price_per_people = price / num_people
        names = dish[2:]
        # iterate over names 
        for name in names:
            # calculate price per person and add to persons total 
            bill_dict[name] = price_per_people+ bill_dict.get(name, 0)
    for name, value in bill_dict.items():
        bill_dict[name] = round(value, 2) # round to 2 
    return bill_dict
#raise NotImplementedError()
```

<!-- grade_id: Q8b-test, points: 10, grade: true, solution: false -->
```python
# Test your function with the provided test input files

assert_equal(split_bill("data/billsplit_input0.txt"), {'Michael': 29.82})
assert_equal(split_bill("data/billsplit_input1.txt"), 
            {'Jamie': 24.31,
             'Victoria': 10.64,
             'Allison': 10.64,
             'Amanda': 2.42,
             'Samuel': 2.42,
             'Peter': 25.88,
             'Lisa': 10.37})

# These tests only verify the length (i.e., number of people); feel free to modify it accordingly.

assert_equal(len(split_bill("data/billsplit_input0.txt")), 1)
assert_equal(len(split_bill("data/billsplit_input1.txt")), 7)
assert_equal(len(split_bill("data/billsplit_input2.txt")), 18)
assert_equal(len(split_bill("data/billsplit_input3.txt")), 45)
assert_equal(len(split_bill("data/billsplit_input4.txt")), 142)

```

## Exercise 9: Dice math (15 points)

In this exercise you will work with .txt files where a number of dice are drawn using characters as follows

```
+-------+ 
|       |
|   O   |
|       |
+-------+

+-------+
| O     |
|       |
|     O |
+-------+

+-------+
|     O |
|       |
| O     |
+-------+

+-------+
| O     |
|   O   |
|     O |
+-------+
        
+-------+
|     O |
|   O   |
| O     |
+-------+
        
+-------+
| O   O |
|       |
| O   O |
+-------+
       
+-------+
| O   O |
|   O   |
| O   O |
+-------+
       
+-------+
| O   O |
| O   O |
| O   O |
+-------+

+-------+
| O O O |
|       |
| O O O |
+-------+
```

In particular note that the value is drawn using the letter ``O``. Note that any such .txt input file might have no dice as well as up to 25 dice. As illustrated above, some of the dice can have different orientations. Multiple dice do *not* necessarily appear in distinct rows of the file (open some of the provided .txt input files to get a feeling), but they never overlap. 

### Part A (5 points)
Write a function that a .txt file containing the representation of a set of dice and returns only the nonnegative integer number equal to the total sum of the values of the dice.

<!-- grade_id: Q9a-code, points: 0, grade: false, solution: true -->
```python
def dice_sum(filename):

    DICE_WIDTH = 9
    DICE_HEIGHT = 5 
    total = 0
    try:
        with open(filename, 'r') as f:
            content = f.read()
            # get count of 'O'
            count = content.count('O')
        return count
    except FileNotFoundError:
        print("File not found.")
        return None 
    
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: Q9a-test, points: 5, grade: true, solution: false -->
```python
# Test your function with the provided test input files

assert_equal(dice_sum("data/dicemath_input0.txt"), 39)
assert_equal(dice_sum("data/dicemath_input1.txt"), 19)
assert_equal(dice_sum("data/dicemath_input2.txt"), 0)
assert_equal(dice_sum("data/dicemath_input3.txt"), 16)

```

### Part B (10 points)

Write a function that reads the .txt file containing the representation of a set of dice and returns only one (possibly negative) integer equal to the difference between the sum of the the even dice minus the sum of the odd dice. 

<!-- grade_id: Q9b-code, points: 0, grade: false, solution: true -->
```python
def dice_diff(filename):
    
    DICE_WIDTH = 9
    DICE_HEIGHT = 5
   
    #YOUR CODE HERE
    with open(filename, 'r') as f:
        content =f.readlines()  
    # get number of rows and columns 
    rows = len(content)
    columns = len(content[0])
    # find conditions 
    val_dict =find_conditions(content, rows, columns, DICE_WIDTH, DICE_HEIGHT)
    # extract dice values 
    dice_values =extract_values(content, val_dict, DICE_WIDTH, DICE_HEIGHT)
    sum_even_values = calculate_sum_even(dice_values)
    sum_odd_values = calculate_sum_odd(dice_values)
    difference = sum_even_values - sum_odd_values
    return difference

def calculate_sum_even(dice_values):
    # calculate sum of even values 
    sum_even_values = 0
    for val in dice_values:
        if val %2 == 0:
            sum_even_values+=val

    return sum_even_values

def calculate_sum_odd(dice_values):
    # calculate sum of odd values 

    sum_odd_values = 0
    for val in dice_values:
        if val %2 != 0:
            sum_odd_values += val
    return sum_odd_values

def find_conditions(content, rows, columns, dice_width, dice_height):
    val_dict = {}
    for i in range(rows):
        for j in range(columns):
            # check if it is corner or not 
            if j < columns -dice_width + 1 and content[i][j] == '+':
                if content[i][j + dice_width -1] == '+' and content[i + 1][j] == '|':
                    val_dict[(i, j)] = 0
    return val_dict

def is_corner(content, i, j, dice_width, dice_height):
    # check if it is corner or not 
    return content[i][j - dice_width + 1] =='+' and content[i - 1][j] == '|'

def extract_values(content, val_dict, dice_width, dice_height):
    dice_values = []
    for i in range(len(content)):
        for j in range(len(content[i])):
            letter = content[i][j]
            # letter 'O' case 
            if letter =='O':
                for key, value in val_dict.items():
                    key_row, key_col = key
                    if key_row <i < (key_row +dice_height - 1) and key_col < j < (key_col + dice_width - 1):
                        # update key's value by incrementing 1 
                        val_dict[key] = value+ 1
                        break
            # letter 'O' case 
            if letter == '+' and dice_width - 1< j:
                if is_corner(content, i, j, dice_width, dice_height):
                    # remove from dictionary 
                    die_value = val_dict.pop((i - dice_height + 1, j - dice_width + 1))
                    dice_values.append(die_value)
    return dice_values
    #raise NotImplementedError()
```

<!-- grade_id: Q9b-test, points: 10, grade: true, solution: false -->
```python
# Test your function with the provided test input files

assert_equal(dice_diff("data/dicemath_input0.txt"), -3)
assert_equal(dice_diff("data/dicemath_input1.txt"), -15)
assert_equal(dice_diff("data/dicemath_input2.txt"), 0)
assert_equal(dice_diff("data/dicemath_input3.txt"), 8)

```

```python

```